Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Viktoriia Stepanenko"
COLLABORATORS = "------------"
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    current_max=0 #the variable to store current sum
    value_max=0 #the variable with the actual max for now
    ind_start=0
    ind_end=0
    for i in range(0, len(A)-2): #i is the index of the first buy date till the last buy date, which would be the second last element of A
        for j in range(i+1, len(A)-1): #j is the index of the sell date
            current_max=A[j]+A[i] #adding the change
            if j>=2:
                for k in range(i+1, j):
                    current_max=current_max+A[k]
            else:
                continue
            if current_max > value_max: #comparing the values
                value_max=current_max
                ind_start=i
                ind_end=j #updating the values
            else:
                continue
    return ind_start, ind_end, value_max
    raise NotImplementedError()
    
bruteforce_max_subarray([-2, -3, 4, -1, -2, 1]) #got confused in my own scheme and not comprehending where I got lost...
```

Output:
```
(2, 3, 3)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
In the case where I can solve the answer analytically - no. With this new array, it is clear that the new subarray will be [1, 8], and the output I would get if I would re-run the algorithm would be (5, 6, 9)

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
True. In the case if we get a new element added in the end of our array, there is no telling what this element will be. It may be the negative element that would be appended, and it would not make sense to include it in the maximum subarray.

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    
    """
    bord=x.index(mx)
    print(x, mx, bord)
    A=x[bord:len(x)]
    print(A)
    current_max=0 #the variable to store current sum
    value_max=0 #the variable with the actual max for now
    ind_start=0
    ind_end=0
    for i in range(0, len(A)-2): #i is the index of the first buy date till the last buy date, which would be the second last element of A
        for j in range(i+1, len(A)-1): #j is the index of the sell date
            current_max=A[j]+A[i] #adding the change
            if j>=2:
                for k in range(i+1, j):
                    current_max=current_max+A[k]
            else:
                continue
            if current_max > value_max: #comparing the values
                value_max=current_max
                ind_start=i
                ind_end=j #updating the values
            else:
                continue
    return ind_start, ind_end, value_max
    raise NotImplementedError()
incremental_max_subarray([-2, -3, 4, -1, -2, 1, 8], 4)

#went somewhere wrong with the addition, will try to turn the list around.
#hope I understood the prompt correctly?
```

Output:
```
[-2, -3, 4, -1, -2, 1, 8] 4 2
[4, -1, -2, 1, 8]
(0, 3, 2)
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

Output:
```
[-2, -3, 4, -1, -2, 1, 8] 4 2
[4, -1, -2, 1, 8]
AssertionError                            Traceback (most recent call last)
<ipython-input-13-2256e7bbc067> in <module>
      1 B = [-2, -3, 4, -1, -2, 1, 8]
----> 2 assert(incremental_max_subarray(B, 4) == 10)
      3 assert(incremental_max_subarray(B[:1], 0) == B[0])

AssertionError:
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
#will redo my code for the previous part once I figure it out.
#also need clarification on this part of the prompt as well. How is it different from a brute force then? How do I define mx? It seemed as though my mx is what I decide to be the start/end of the list?

def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
This approach fully depends on whether or not you know exactly where the maximum subarray starts/ends, so no, the algorithm that decides by itself is certainly more efficient than that.

```python

```