Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **50 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Viktoriia Stepanenko"
COLLABORATORS = "----"
```

---

# CS110 Pre-class Work - Analyzing elementary sorting algorithms



## Question 1 [time estimate: 10 minutes]
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
This line ensures that when we are moving the element in the attempts to find its place, while the comparison loop runs, we examine, say, second and third element on the left as well, not just the first one. By changing the index and making it smaller, we move to the left and check whether our current key element (e.g. with previous index 3) is smaller than the element with index 2, and if it is, then we move A[2] on the place where our key initially was and check the element with index 1 (if our key is smaller than this one too, then we move this element to the right as well). Thus, this line ensures that we check all the relevant elements while the loop runs.

## Question 2 [time estimate: 15 minutes]
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertion_sort(A):
    for j in range(2, len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertion_sort(A):
    for j in range(1, len(A)): #!this loop goes through all the elements of the list correctly now!
        
        #THE BUG: the loop was starting from 3nd element of the list instead of 2nd, although we know that in Python the count of indeces starts from 0, and not 1. Thus, to start this loop correctly, we need to set the starting point of the range to 1 (so that we start with comparing the 2nd element with the first)
        
        key = A[j] #store the value of the element we're testing
        print('key is', key)
        i = j-1 #store the index of the element on the left from where the element we're testing is initially located | the index of the element we will be comparing to
        while i >= 0 and A[i]>key: #1. till we reach the start of the list in our comparisons; 2. till the element on the left is larger than our element -> keep reshuffling the elements
            A[i+1] = A[i] #move the element that was on the left to our initial position (or the position that belonged to the element of A[i+1] in the moment of the restart of the loop)
            i -= 1 #move to the element on the left for the examination
        A[i+1] = key #put the element we were searching the place for on the place that was "freed" for it by shuffling the elements from the left that were larger than it to the right
        print(A) #for me to check out how the algorithm shuffles the list
    return A #return sorted list
    raise NotImplementedError()

insertion_sort([10,2,7,5,3,6]) #tested and got the correct result. Before fixing the bug, all the elements were returned sorted, besides 10.
```

Output:
```
key is 2
i is 0
A[i+1] is 2
A[i] is 10
[10, 10, 7, 5, 3, 6]
i on the 2nd step is -1
[2, 10, 7, 5, 3, 6]
key is 7
i is 1
A[i+1] is 7
A[i] is 10
[2, 10, 10, 5, 3, 6]
i on the 2nd step is 0
[2, 7, 10, 5, 3, 6]
key is 5
i is 2
A[i+1] is 5
A[i] is 10
[2, 7, 10, 10, 3, 6]
i on the 2nd step is 1
A[i+1] is 10
A[i] is 7
[2, 7, 7, 10, 3, 6]
i on the 2nd step is 0
[2, 5, 7, 10, 3, 6]
key is 3
i is 3
A[i+1] is 3
A[i] is 10
[2, 5, 7, 10, 10, 6]
i on the 2nd step is 2
A[i+1] is 10
A[i] is 7
[2, 5, 7, 7, 10, 6]
i on the 2nd step is 1
A[i+1] is 7
A[i] is 5
[2, 5, 5, 7, 10, 6]
i on the 2nd step is 0
[2, 3, 5, 7, 10, 6]
key is 6
i is 4
A[i+1] is 6
A[i] is 10
[2, 3, 5, 7, 10, 10]
i on the 2nd step is 3
A[i+1] is 10
A[i] is 7
[2, 3, 5, 7, 7, 10]
i on the 2nd step is 2
[2, 3, 5, 6, 7, 10]
[2, 3, 5, 6, 7, 10]
```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertion_sort([0]) == [0])
assert(insertion_sort([-1,1]) == [-1,1])
assert(insertion_sort([1,-1]) == [-1,1])
assert(insertion_sort([1,6,3,6]) == [1,3,6,6])
assert(insertion_sort([7, 3, 5, 4, 6, 4]) == [3, 4, 4, 5, 6, 7])
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
Besides me trying things out myself and seeing the algorithm working - the assert() function returns nothing, meaning the condition (that the input into the function produces the expected output) is True. If that would not be the case, it would raise an Assertion error.

## Question 3 [time estimate: 15 minutes]
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertion_sort`,  your new Python function, `bubble_sort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
#pseudo-code:
#BUBBLESORT(A)
#for i=1 to A.length-1
#for j=A.length downto i+1
#if A[j]<A[j-1]
#exchange A[j] with A[j-1]

def bubble_sort(A):
    for i in range(0, len(A)-1): #!I have changed the i as a starting point from 1 to 0 here :/ Unsure if that was meant to be a bug or not, but it worked better for me in practice, so I left it be and highlighted for myself to double-check and ask about later. If I would leave 1 to be, then my code would not account for the first element of the list at all. Now this code goes through all of them.!
        for j in reversed(range(i+1, len(A))): #j starts the count from the end of the list
            if A[j]<A[j-1]: #if our element is smaller than the element on the left - swap their places
                key=A[j]
                A[j]=A[j-1]
                A[j-1]=key
                print(A) #for me to check out how the algorithm shuffles the list
            else:
                continue
    return A
    raise NotImplementedError()

assert(bubble_sort([2,5,6,1,9,4]) == [1,2,4,5,6,9])
assert(bubble_sort([8,3,1,7,5]) == [1,3,5,7,8])
assert(bubble_sort([10,90,30,100]) == [10,30,90,100])
```

Output:
```
[2, 5, 6, 1, 4, 9]
[2, 5, 1, 6, 4, 9]
[2, 1, 5, 6, 4, 9]
[1, 2, 5, 6, 4, 9]
[1, 2, 5, 4, 6, 9]
[1, 2, 4, 5, 6, 9]
[8, 3, 1, 5, 7]
[8, 1, 3, 5, 7]
[1, 8, 3, 5, 7]
[1, 3, 8, 5, 7]
[1, 3, 5, 8, 7]
[1, 3, 5, 7, 8]
[10, 30, 90, 100]
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
You can see my test cases above and see that the code passes all of them.
You can say my program is bug-free, since it solves the problem I want it to solve correctly, but whether it is correct depends on the definition of correct. Since I have made my own adjustments and there is no guarantee that my code perfectly implements intended pseudo-code (I may have misinterpreted the instructions), then there is a question of whether it is correct in the sense of: "Does it produce the intended algorithm correctly?". From seeing the way my algorithm works, I hope that yes.

## Question 4 [time estimate: 10 minutes]
The following pseudocode corresponds to a third sorting algorithm covered in this lesson, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free by providing at least three test cases. `selection_sort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
#SELECTION-SORT(A)
#n=A.length
#for i=1 to n-1
#minidx = i
#for j=i+1 to n
#if A[j]<A[minidx] then
#minidx = j
#swap A[i] with A[minidx]

def selection_sort(A):
    n=len(A)
    for i in range(0, n-1): #changed the starting point once again, for the same purposes as before (the algorithm starts actually fulfilling its purpose) :/
        minidx=i #index of the first element of the unsorted part of the row and the "min" that we take just as a guess and later compare to other numbers and reevaluate which one is the real min
        for j in range(i+1, n): #go through all the other elements to check whether we need to change our judgement on what is the real min in this try
            if A[j]<A[minidx]: #if the element we are checking in this try is smaller than what we thought was min, reassign what the min is
                minidx=j
            else:
                continue
        key=A[minidx] #once we have determined the real min, swap the places of this min with our previous guess
        A[minidx]=A[i]
        A[i]=key
        print(A) #for me to check out how the algorithm shuffles the list

    return A
    raise NotImplementedError()

assert(selection_sort([2,5,6,1,9,4]) == [1,2,4,5,6,9]) #all the tests seem to work fine
assert(selection_sort([8,3,1,7,5]) == [1,3,5,7,8])
assert(selection_sort([10,90,30,100]) == [10,30,90,100])
```

Output:
```
[1, 5, 6, 2, 9, 4]
[1, 2, 6, 5, 9, 4]
[1, 2, 4, 5, 9, 6]
[1, 2, 4, 5, 9, 6]
[1, 2, 4, 5, 6, 9]
[1, 3, 8, 7, 5]
[1, 3, 8, 7, 5]
[1, 3, 5, 7, 8]
[1, 3, 5, 7, 8]
[10, 90, 30, 100]
[10, 30, 90, 100]
[10, 30, 90, 100]
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**